```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

<div class="alert alert-success">

#### Homework 1 Supplemental Notebook
    
# Means, Sums, and Calculus

### EECS 245, Fall 2025 at the University of Michigan
    
</div>

### Instructions

Most homeworks will have Jupyter Notebooks, like this one, designed to supplement the theoretical problems. 

To write and run code in this notebook, you have two options:

1. **Use the EECS 245 DataHub.** To do this, click the link provided in the Homework 1 PDF. Before doing so, read the instructions on the [**Tech Support**](https://eecs245.org/tech-support/#option-1-using-the-eecs-245-datahub) page on how to use the DataHub.
1. **Set up a Jupyter Notebook environment locally, and use `git` to clone our course repository.** For instructions on how to do this, see the [**Tech Support**](https://eecs245.org/tech-support) page of the course website.

To receive credit for the programming portion of the homework, you'll need to submit your completed notebook to the autograder on Gradescope. Your submission time for Homework 1 is the **latter** of your PDF and code submission times.

Remember that homework problems have hidden test cases. The public test cases in your notebook only verify that your answer is in the correct format and on the right track; your results on the hidden tests will be available to you on Gradescope after we release grades.

```python
# Run this cell before proceeding; it's needed for the autograder to work.
import numpy as np
```

## Problem 6: Fun with Loops (8 pts)

---

In this problem, you'll work through a few separate tasks that involve writing loops in Python, both `for`-loops and `while`-loops.

But first, for some context, since we didn't discuss loops in Lab 1. In Python, you can loop over any **iterable**. Strings, lists, and dictionaries are all examples of iterables.

All of the following are valid ways to write a `for`-loop.<br>
    
```py
            for value in "this is a string":

            for element in lst:                  # Assume lst is a list.

            for i in range(len(lst)):
    
```

```python
sentence = ['the', 'university', 'of', 'michigan', 'at', 'ann', 'arbor']

# word is 'the', then 'university', then 'of', ...
for word in sentence:
    print(word)
```

```python
# i is 0, then 1, then 2, ...
# Depending on the problem, you may need to use the first pattern or this second pattern.
for i in range(len(sentence)):
    print(i, sentence[i])
```

One of the more common `for`-loop examples you may have seen in earlier classes involved performing some operation to every element of a sequence, e.g. doubling the numbers in a list.

```py
def double(vals):
    new_vals = []
    for val in vals:
        new_vals.append(vals * 2)
    return new_vals
```

**Warning**: We are going to **avoid ❌** these kinds of `for`-loops in this class, because there are **much faster** ways of achieving the same goal in `numpy`, as we'll learn about in Lab 2.

`while`-loops exist too, but they'll come up sparingly in the context of machine learning.

### Problem 6a) (2 pts)

In math, the word "alternating" is sometimes used to indicate that the signs of the elements of a sequence oscillate, or "bounce", back and forth between positive and negative.

Complete the implementation of the function `alternating_sign_mean`, which takes in a list of positive numbers, `values`, and computes the mean of **_every_ element** in the original list with **alternating signs**, starting with a positive sign for the element at position 0, a negative sign for the element at position 1, and so on. Example behavior is shown below.

```python
>>> alternating_sign_mean([3, 2, 9.5])
3.5 # Comes from (3 + (-2) + 9.5) / 3.

>>> alternating_sign_mean([3, 2.5, 7, 6.5])
0.25 # Comes from (3 + (-2.5) + 7 + (-6.5)) / 4.
```

Note that you **should not** use any additional packages (e.g. `math`, `numpy`) in your solution – everything should be done using base Python.

```python
import numpy as np

def alternating_sign_mean(values):
    ...
    
# Feel free to change this input to make sure your function works correctly.
alternating_sign_mean([3, 2, 9.5])
```

```python
grader.check("p06_a")
```

### Problem 6b) (3 pts)

Suppose we are waiters at a restaurant, and we keep track of the tips that we receive in a given night. Specifically, we store in a list the first names of our customers along with the tip that they gave, like so:

```python
[name_1, tip_1, name_2, tip_2, name_3, tip_3, ...]
```

Names are strings, while tips are non-negative numbers, either `int` or `float`.

Complete the implementation of the function `big_tippers`, which takes in a list `tips` as described above and returns only the names of the customers who tipped **strictly more** than the average tip. Example behavior is given below.

```python
>>> tips = ['angela', 15, 'vincent', 9, 'pranavi', 19]
>>> big_tippers(tips)
['angela', 'pranavi']
```

In the example above, the average tip is $\frac{15 + 9 + 19}{3} = 14.33$, and both `'angela'` and `'pranavi'` tipped more than that, so your function should return `['angela', 'pranavi']`.

You may assume that `tips` is well-formed, i.e. has an even, positive number of entries, in the correct order. Also assume that customer names are unique, i.e. that there are no duplicate tippers.

Note that you **should not** use any additional packages (e.g. `math`, `numpy`) in your solution – everything should be done using base Python.

```python
def big_tippers(tips):
    ...

# Feel free to change this input to make sure your function works correctly.
big_tippers(['angela', 15, 'vincent', 9, 'pranavi', 19])
```

```python
grader.check("p06_b")
```

### Problem 6c) (1.5 pts)

Suppose `vals` is a list of length $n-1$, containing the integers from 1 to $n$, inclusive, with no duplicates and in unsorted order, but with exactly 1 value missing. Complete the implementation of the function `missing_value_naive`, which takes in a list like `vals` and returns the number that is missing, as an `int`. Example behavior is given below.

```py
>>> missing_value_naive([2, 1, 5, 3])
4

>>> missing_value_naive([10, 9, 1, 6, 5, 7, 2, 4, 3])
8
```

```python
def missing_value_naive(vals):
    ...

# Feel free to change this input to make sure your function works correctly.
missing_value_naive([2, 1, 5, 3])
```

```python
grader.check("p06_c")
```

### Problem 6d) (1.5 pts)

Your implementation of `missing_value_naive` in Problem 6c) didn't need to be particularly efficient, and you likely used a `for`-loop. A `for`-loop that involves a call to the `in` operator on each call results in an $O(n^2)$ implementation. (Don't worry if you're not familiar with $O$-notation – all this means is that the speed of your implementation grows very quickly as the length of the input list grows, which is not a good sign if we care about the efficiency of our code.)

To make our solution to the missing value problem more efficient, we’ll use the fact that the missing value is the difference between the sum we’d expect if none of the values were missing, and the sum
of the values we actually have. For example, if `vals = [1, 2, 5, 3]`, then $n = 5$, so the sum we’d
expect is $1+2+3+4+5 = 15$, and so the missing value is $(1+2+3+4+5)−(1+2+5+3) = 15−11 = 4$.

Conveniently, in Lab 1, we provided you with the following fact:

$$\sum_{i=1}^n i = 1 + 2 + \ldots + n = \frac{n(n+1)}{2}$$

This means that to find the missing value, we can compute the difference between the sum if none were missing – using the formula above – and the sum of the values actually present. This will enable us to find the missing value with just a single pass over the list (to compute the sum of the present values), leading to an $O(n)$ solution.

Complete the implementation of `missing_value_fast`, which uses the logic presented here to find the missing value in `vals`.

- Make sure that `missing_value_fast`, like `missing_value_naive`, returns an `int`.

- **Do not** use `sum(range(1, n+1))` in your implementation – your implementation can only use the `sum` function once.

```python
def missing_value_fast(vals):
    ...
    
# Feel free to change this input to make sure your function works correctly.
missing_value_fast([2, 1, 5, 3])
```

```python
grader.check("p06_d")
```

To test whether your implementation of `missing_value_fast` is significantly faster than your implementation of `missing_value_naive`, we will use the function `relative_speed`, defined below.

```python
import numpy as np
import time
def relative_speed(reps=100, n=1000):
    t1 = time.time()
    for _ in range(reps):
        missing_value_naive(np.random.permutation(range(1, n+1))[:-1])
    delta1 = time.time() - t1
    
    t2 = time.time()
    for _ in range(reps):
        missing_value_fast(np.random.permutation(range(1, n+1))[:-1])
    delta2 = time.time() - t2
    
    return delta1 / delta2
```

When you call `relative_speed` on the default arguments, you should see a speedup of at least 6x on average (though the Gradescope autograder will be slightly more lenient).

```python
relative_speed()
```

## Finish Line 🏁

Congratulations! You're ready to submit the programming portion of Homework 1.

To submit your work to Gradescope:

1. Select `Kernel -> Restart & Run All` to ensure that you have executed all cells, including the test cells.
2. Read through the notebook to make sure everything is fine and all public tests passed.
3. Run the cell below to run all tests, and make sure that they all pass.
4. Download your notebook using `File -> Download`, then upload your notebook to Gradescope under "Homework 1, Problem 6 Code".
5. Stick around for a few minutes while the Gradescope autograder grades your work. Make sure you see that all **public tests** have passed on Gradescope. **Remember that homeworks have hidden tests!**

Tests for p06_a (points: 2)
```
>>> np.isclose(alternating_sign_mean([3, 1.5, 9.9]), 3.8)
True
```
```
>>> np.isclose(alternating_sign_mean([2, 4.5, 3, 7]), -6.5 / 4)
True
```


Tests for p06_b (points: 3)
```
>>> big_tippers(['angela', 15, 'vincent', 9, 'pranavi', 19]) == ['angela', 'pranavi']
True
```


Tests for p06_c (points: 1.5)
```
>>> callable(missing_value_naive) and isinstance(missing_value_naive([2, 1, 5, 3]), int)
True
```
```
>>> ex1 = missing_value_naive([2, 1, 5, 3])
>>> ex2 = missing_value_naive([10, 9, 1, 6, 5, 7, 2, 4, 3])
>>> ex1 == 4 and ex2 == 8
True
```


Tests for p06_d (points: 1.5)
```
>>> callable(missing_value_fast) and isinstance(missing_value_fast([2, 1, 5, 3]), int)
True
```
```
>>> ex1 = missing_value_fast([2, 1, 5, 3])
>>> ex2 = missing_value_fast([10, 9, 1, 6, 5, 7, 2, 4, 3])
>>> ex1 == 4 and ex2 == 8
True
```
